# CSCI 3155: Assignment 9

Topics: 
- Type Inference
- Basics of Objects
- Variance Annotations.

Readings: Notes posted on moodle.  Chapters from Odersky

- Chapter 4 (revise)
- Chapter 10
- Chapter 11 (take note of Any, AnyVal, Null, Nothing)
- Chapter 12 
- Chapter 19

__Name__: WRITE YOUR NAME HERE

```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (20 points): Objects in Scala

### A (12 points)
Consider the class below for an employee record in an organization that has three fields `name`, `eid` (employee id) and `salary`.

1) As given, scala will complain that class employee cannot extend the abstract class People. Add the missing methods to enable the class Employee to properly extend People

2) Write a factory pattern that takes in a formatted string of the form "John Doe,29585,33110" consisting of a string and two numbers, to extract the name and salaries for each employee. You may want to use `scala.util.matching.Regex` [https://www.scala-lang.org/api/2.9.2/scala/util/matching/Regex.html]
If the input fails to match the pattern, an `IllegalArgumentException` must be thrown

Companion objects are covered in scala book Chapter 4.3. Also lookup here: https://alvinalexander.com/scala/factory-pattern-in-scala-design-patterns

<!-- grade_id: cell-f80a491b407a7ab7, points: 0, grade: false, solution: true -->
```scala
import scala.util.matching.Regex

abstract class People {
    /* I need the ability to convert every person to a string */
    def toString: String
    /* I need the ability to check if two people are the same */
    def equals (p: People): Boolean
}

/* 
  Note that for an employee to be equal to an object of type People, that object
  must be an employee with matching eid number 
*/

class Employee( val name: String, val eid: Int, var salary: Int) extends People {
    //BEGIN SOLUTION
    override def toString: String = {
       (s"Employee Name: $name, EID: $eid, Salary: $salary")
    }
    
    override def equals(p: People): Boolean = {
        if (p.isInstanceOf[Employee]){
            val q = p.asInstanceOf[Employee]
            q.eid == this.eid
        } else 
            false
        
    }
    //END SOLUTION
}

object Employee {
    def apply(formattedInput: String): Employee = {
        /* 
        Input format must be 
        (full name with possible spaces)[optional whitespaces],[optional whitespaces](employee ID number)[optional whitespaces],[optional whitespaces](employee salary)
        */
        //BEGIN SOLUTION
        val rex = raw"([A-Za-z ]*)\s*,\s*(\d+)\s*,\s*(\d+)".r
        formattedInput match {
            case rex(name, i1, i2) => { new Employee(name.trim, i1.toInt, i2.toInt)}
            case _ => throw new IllegalArgumentException(s"Badly formatted input $formattedInput")
        }
        //END SOLUTION
    }
}
```

Output:
```
[32mimport [39m[36mscala.util.matching.Regex

[39m
defined [32mclass[39m [36mPeople[39m
defined [32mclass[39m [36mEmployee[39m
defined [32mobject[39m [36mEmployee[39m
```

<!-- grade_id: cell-e1ec3bf75d6f6de5, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
println(e1.toString)
assert(e1.name == "Jonathan Doe", "Failed")
assert(e1.eid == 12893, "Failed")
assert(e1.salary == 39813, "Failed")
assert(e1.equals(new Employee("whoknows", 12893, 39813)))
passed(4)
//END TEST
```

Output:
```
Employee Name: Jonathan Doe, EID: 12893, Salary: 39813

*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Employee Name: Jonathan Doe, EID: 12893, Salary: 39813
```

<!-- grade_id: cell-9d823f9842fd95a7, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
try {
    val e2 = Employee("Badly Formatted String, 29812, $3102")
    assert(false, "Failed, the string is badly formatted but your code is OK with it.")
} catch {
    case e => println(s"Expected behavior seen")
}
passed(4)
//END TEST
```

Output:
```
Expected behavior seen

*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-00d1907d97c7da69, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Janet Maria Catherine Doe    , 12894   , 121331")
println(e1.toString)
assert(e1.name == "Janet Maria Catherine Doe", "Failed")
assert(e1.eid == 12894, "Failed")
assert(e1.salary == 121331, "Failed")
assert(e1.equals(new Employee("random", 12894, 121332)))
passed(4)
//END TEST
```

Output:
```
Employee Name: Janet Maria Catherine Doe, EID: 12894, Salary: 121331

*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Employee Name: Janet Maria Catherine Doe, EID: 12894, Salary: 121331
```

### B (8 points)
We wish to add a new class of people to our system. These are "Contractor" with the information associated should involve their name (string), contract ID (Integer) and contract supervisor (Employee).

Write a new class named `Contractor` whose constructor should take in the name of the person (String), their ID number (Integer) and Supervisor (Employee) and extend from People. 

A contractor object is equal to an object p of type People, only if p is actually an instance of contractor, and if their names are the same and they have the same contract ID number.

<!-- grade_id: cell-97ab958367cfbd0a, points: 0, grade: false, solution: true -->
```scala
//BEGIN SOLUTION
class Contractor (val name: String, val idNumber: Int, val e: Employee) extends People {
    override def toString: String = s"Contractor $name, id: $idNumber, supervisor: $e"
    override def equals(p: People) = {
        if (p.isInstanceOf[Contractor]){
            val q = p.asInstanceOf[Contractor]
            (q.name == name) && (q.idNumber == idNumber)
        } else false
    }
}
//END SOLUTION
```

Output:
```
defined [32mclass[39m [36mContractor[39m
```

<!-- grade_id: cell-82cc00e175a12efe, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
val c1 = new Contractor("John Contract Worker", 12093, e1)
assert( !c1.equals(e1), "Failed")
assert( c1.equals(c1), "Failed")
passed(4)
//END TEST
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Employee Name: Jonathan Doe, EID: 12893, Salary: 39813
[36mc1[39m: [32mContractor[39m = Contractor John Contract Worker, id: 12093, supervisor: Employee Name: Jonathan Doe, EID: 12893, Salary: 39813
```

<!-- grade_id: cell-62fa2e0b66f58862, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
val e2 = Employee("Jane Doe, 12894, 43814")
val c1 = new Contractor("John Contract Worker", 12093, e2)
val c2 = new Contractor("John Contract Worker", 12093, e1)
assert( c1.equals(c2), "Failed")
passed(4)
//END TEST
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Employee Name: Jonathan Doe, EID: 12893, Salary: 39813
[36me2[39m: [32mEmployee[39m = Employee Name: Jane Doe, EID: 12894, Salary: 43814
[36mc1[39m: [32mContractor[39m = Contractor John Contract Worker, id: 12093, supervisor: Employee Name: Jane Doe, EID: 12894, Salary: 43814
[36mc2[39m: [32mContractor[39m = Contractor John Contract Worker, id: 12093, supervisor: Employee Name: Jonathan Doe, EID: 12893, Salary: 39813
```

## Problem 2 (40 points): Traits in Scala

For this problem, we have defined two traits: `NumberOfLegs` that helps us define how many legs a given animal has and `WarmBlooded` that applies to warm blooded animals.

```scala
abstract class Animal 

trait NumberOfLegs {
    val nLegs: Int
    def getNumberOfLegs: Int = nLegs
}

trait WarmBlooded extends Animal {
    val bodyTempMaintained: Double
    def getBodyTemp: Double = bodyTempMaintained 
}
```

Output:
```
defined [32mclass[39m [36mAnimal[39m
defined [32mtrait[39m [36mNumberOfLegs[39m
defined [32mtrait[39m [36mWarmBlooded[39m
```

### A (5 points)

Define a class `Human` that inherits from `Animal` and mixes in the traits `NumberOfLegs` with `nLegs = 2` and `WarmBlooded` with `bodyTempMaintained = 98`. The class `Human` must take in a parameter called `name` of type `String` and implement a `getName` method without any parameters.

<!-- grade_id: cell-e7a4a9275bd24e24, points: 0, grade: false, solution: true -->
```scala
//BEGIN SOLUTION
class Human(val name: String) extends Animal with NumberOfLegs with WarmBlooded {
    override val nLegs:Int = 2
    val bodyTempMaintained = 98.0
    def getName: String = name
}
//END SOLUTION
```

Output:
```
defined [32mclass[39m [36mHuman[39m
```

<!-- grade_id: cell-f9b7ae8902326372, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val t1 = new Human("Jane Smith")
assert(t1.getNumberOfLegs == 2, "Your human does not have two legs")
assert(t1.bodyTempMaintained == 98.0, "Your human does not maintain a body temp of 98")
assert(t1.getBodyTemp == 98.0, "Your human's getBodyTemp Function is not working")
assert(t1.getName == "Jane Smith", "Your human's name is not setting correctly")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36mt1[39m: [32mHuman[39m = ammonite.$sess.cmd10$Helper$Human@8cc013c
```

### B (5 points)
Define a class named `Table` that mixes in the trait `NumberOfLegs` with `nLegs = 4`. 

_Note:_ `Table` cannot mixin the trait `WarmBlooded` (make sure you understand why that is as a quick check of concepts).

<!-- grade_id: cell-8145df690b6f6588, points: 0, grade: false, solution: true -->
```scala
//BEGIN SOLUTION
class Table extends NumberOfLegs {
    val nLegs = 4
}
//END SOLUTION
```

Output:
```
defined [32mclass[39m [36mTable[39m
```

<!-- grade_id: cell-d34b29b062b93894, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val tbl = new Table()
assert(tbl.getNumberOfLegs == 4, "A Table must have four legs")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36mtbl[39m: [32mTable[39m = ammonite.$sess.cmd12$Helper$Table@2c811329
```

### C (15 points)
Define a class called `Menagerie` with type parameter `T` such that 

- T must always be a derived type of Animal with the trait NumberOfLegs (In scala you can say `T <: A with B` to indicate that a type parameter T must derive from class A with trait B applied).

- `Menagerie` has a list of objects of type `T`.

- The user must be able to instantiate an object of type Menagerie without providing any class parameters: `new Menagerie[T]`

- The class Menagerie has the method `addAnimalToMenagerie` that adds an object of type `T` to menagerie and returns a new Menagerie.

- Implement a method `addMultipleAnimalsToMenagerie` that adds more than one animal. Lookup syntax below.
 
 
Lookup the syntax for varargs or repeated parameters (T*) here: https://alvinalexander.com/scala/how-to-define-methods-variable-arguments-varargs-fields. Also in Chapter 8.8 of Odersky book.


- Add a method `totalLegs: Int` that sums up how many legs the animals in the menagerie have in total.

~~~
class Menagerie[T...] () {
   // ... Extra Stuff that you may want to add ...
   def addAnimalToMenagerie (a: T) : Menagerie[T]
   def addMultipleAnimalsToMenagerie(animals: T*): Menagerie[T]
   def totalLegs:Int
}
~~~


<!-- grade_id: cell-44221c4360112bc9, points: 0, grade: false, solution: true -->
```scala
//BEGIN SOLUTION
class Menagerie[T <: Animal with NumberOfLegs] (val lst:List[T]) {
    def this() = this(Nil)
    
    def addAnimalToMenagerie(a: T): Menagerie[T] = {
        new Menagerie(a::lst)
    }
    
    def addMultipleAnimalsToMenagerie(animals: T*): Menagerie[T] = {
        
        //val nList: List[T] = animals.foldLeft[List[T]] (lst) { case (lst, a) => a::lst}
        var nList: List[T] = lst
        for (x <- animals) {
            nList = x::nList
        }\
        
        new Menagerie(nList
    }
    
    def totalLegs: Int = lst.map(_.getNumberOfLegs).sum
}
//END SOLUTION
```

Output:
```
defined [32mclass[39m [36mMenagerie[39m
```

```scala
//BEGIN TEST SETUP
class Donkey extends Animal with NumberOfLegs with WarmBlooded {
    val nLegs:Int = 4
    val bodyTempMaintained = 96.0
}

class Lion extends Animal with NumberOfLegs {
    val nLegs:Int = 4
}

class Mule extends Donkey

val d1 = new Donkey()
val d2 = new Donkey()
val d3 = new Donkey()
val d4 = new Mule()
//END TEST SETUP
```

Output:
```
defined [32mclass[39m [36mDonkey[39m
defined [32mclass[39m [36mLion[39m
defined [32mclass[39m [36mMule[39m
[36md1[39m: [32mDonkey[39m = ammonite.$sess.cmd15$Helper$Donkey@7a7420d9
[36md2[39m: [32mDonkey[39m = ammonite.$sess.cmd15$Helper$Donkey@353f9f2d
[36md3[39m: [32mDonkey[39m = ammonite.$sess.cmd15$Helper$Donkey@10b9ff66
[36md4[39m: [32mMule[39m = ammonite.$sess.cmd15$Helper$Mule@6e02182c
```

<!-- grade_id: cell-d70b45ba2426c6a6, points: 3, grade: true, solution: false -->
```scala
//BEGIN TEST
val m1 = new Menagerie[Donkey]
passed(3)
//END TEST
```

Output:
```
java.lang.NullPointerException
  ammonite.$sess.cmd14$Helper$Menagerie.<init>(cmd14.sc:1)
  ammonite.$sess.cmd14$Helper$Menagerie.<init>(cmd14.sc:2)
  ammonite.$sess.cmd16$Helper.<init>(cmd16.sc:1)
  ammonite.$sess.cmd16$.<init>(cmd16.sc:7)
  ammonite.$sess.cmd16$.<clinit>(cmd16.sc:-1)
```

<!-- grade_id: cell-0f3cd5e1e161c390, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST
val m1 = new Menagerie[Donkey]
val m2 = m1.addAnimalToMenagerie(d1)
val n = m2.totalLegs
assert(n == 4, s"1 quadraped = 4 legs. You have $n legs : are you missing legs? Too many legs?")
passed(6)
//END TEST
```

Output:
```
java.lang.NullPointerException
  ammonite.$sess.cmd14$Helper$Menagerie.<init>(cmd14.sc:1)
  ammonite.$sess.cmd14$Helper$Menagerie.<init>(cmd14.sc:2)
  ammonite.$sess.cmd17$Helper.<init>(cmd17.sc:1)
  ammonite.$sess.cmd17$.<init>(cmd17.sc:7)
  ammonite.$sess.cmd17$.<clinit>(cmd17.sc:-1)
```

<!-- grade_id: cell-9ca767949c491818, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST
val m1 = new Menagerie[Donkey]
val m2 = m1.addMultipleAnimalsToMenagerie(d1, d2, d3, d4)
val n = m2.totalLegs
assert(n == 16, s"4 quadrapeds = 16 legs. You have $n legs : are you missing legs? Too many legs?")
passed(6)
//END TEST
```

Output:
```
java.lang.NullPointerException
  ammonite.$sess.cmd14$Helper$Menagerie.<init>(cmd14.sc:1)
  ammonite.$sess.cmd14$Helper$Menagerie.<init>(cmd14.sc:2)
  ammonite.$sess.cmd18$Helper.<init>(cmd18.sc:1)
  ammonite.$sess.cmd18$.<init>(cmd18.sc:7)
  ammonite.$sess.cmd18$.<clinit>(cmd18.sc:-1)
```

### D (15 points)
Currently a `Menagerie` as defined above is Invariant. Make it covariant so that  subtyping relations such as `Menagerie[Mule] <: Menagerie[Donkey] <: Menagerie[Animal]`.

<!-- grade_id: cell-3dcb8201b939b421, points: 0, grade: false, solution: true -->
```scala
//BEGIN SOLUTION
class Menagerie[+T <: Animal with NumberOfLegs] (val lst:List[T]) {
    def this() = this(Nil)
    
    def addAnimalToMenagerie[U >: T <: Animal with NumberOfLegs ](a: U): Menagerie[U] = {
        new Menagerie(a::lst)
    }
    
    def addMultipleAnimalsToMenagerie[U >: T <: Animal with NumberOfLegs ](animals: U*): Menagerie[U] = {
        
        //val nList: List[T] = animals.foldLeft[List[T]] (lst) { case (lst, a) => a::lst}
        var nList: List[U] = lst
        for (x <- animals) {
            nList = x::nList
        }
        new Menagerie(nList)
    }
    
    def totalLegs: Int = lst.map(_.getNumberOfLegs).sum
}
//END SOLUTION
```

Output:
```
defined [32mclass[39m [36mMenagerie[39m
```

<!-- grade_id: cell-96824e578d275d9c, points: 15, grade: true, solution: false -->
```scala
//BEGIN TEST 
class Donkey extends Animal with NumberOfLegs with WarmBlooded {
    val nLegs:Int = 4
    val bodyTempMaintained = 96.0
}

class Mule extends Donkey

val d1 = new Donkey()
val d2 = new Donkey()
val d3 = new Donkey()
val mu1 = new Mule()
val m1 = new Menagerie[Donkey]
// Inheritance among Menageries:
val m2: Menagerie[Donkey] = new Menagerie[Mule]
// Strongest type selection:
val m3: Menagerie[Donkey] = m1.addMultipleAnimalsToMenagerie(d1, d2, d3, mu1)

print("Successfully added Lions and Donkeys to the same menagerie!")

passed(15)
//END TEST
```

Output:
```
java.lang.NullPointerException
  ammonite.$sess.cmd19$Helper$Menagerie.<init>(cmd19.sc:1)
  ammonite.$sess.cmd19$Helper$Menagerie.<init>(cmd19.sc:2)
  ammonite.$sess.cmd20$Helper.<init>(cmd20.sc:12)
  ammonite.$sess.cmd20$.<init>(cmd20.sc:7)
  ammonite.$sess.cmd20$.<clinit>(cmd20.sc:-1)
```

```scala

```